rule_files:
  - ../../resources/prometheus-rules/argocd-controller-rules.yaml

evaluation_interval: 1m

tests:
  # TEST 1: Alert SHOULD fire
  - interval: 1m
    input_series:
      - series: 'argocd_app_reconcile_bucket{le="10"}'
        values: '0+0x5 0+0.1x35'  # Very slow increase for lowest bucket
      - series: 'argocd_app_reconcile_bucket{le="30"}'
        values: '1+0x5 1+0.1x35'  # Slow increase for 30s bucket
      - series: 'argocd_app_reconcile_bucket{le="60"}'
        values: '5+0x5 5+0.5x35'  # Faster increase
      - series: 'argocd_app_reconcile_bucket{le="+Inf"}'
        values: '10+0x5 10+1x35'  # Fastest increase for total

    alert_rule_test:
      - eval_time: 30m  # Must exceed 15m 'for' + 5m rate window + buffer
        alertname: ArgoCDControllerHighReconciliationTime
        exp_alerts:
          - exp_labels:
              severity: warning
            exp_annotations:
              summary: "ArgoCD controller high reconciliation time"
              description: "90% of application reconciliation operations are taking more than 30 seconds to complete"

  # TEST 2: Alert SHOULD NOT fire (90th percentile ≤30s)
  - interval: 1m
    input_series:
      - series: 'argocd_app_reconcile_bucket{le="10"}'
        values: '5+0x30'
      - series: 'argocd_app_reconcile_bucket{le="30"}'
        values: '9+0x30'
      - series: 'argocd_app_reconcile_bucket{le="+Inf"}'
        values: '10+0x30'

    alert_rule_test:
      - eval_time: 20m
        alertname: ArgoCDControllerHighReconciliationTime
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: 'argocd_git_request_total{status_code="400"}'
        values: '0+0x10 0+1x30'
  
    # TEST ArgoCD Repo Server Errors
    alert_rule_test:
      - eval_time: 5m
        alertname: ArgoCDRepoServerRequestErrors
        exp_alerts: []

      - eval_time: 30m
        alertname: ArgoCDRepoServerRequestErrors
        exp_alerts:
          - exp_labels:
              severity: warning
              status_code: 400
            exp_annotations:
              summary: "ArgoCD repository server errors"
              description: "There are Git request errors in ArgoCD repo server"

  # TEST ArgoCD Redis Down
  - interval: 1m
    input_series:
      # No series defined for 'up{job="argocd-redis"}' to test absence

    alert_rule_test:
      - eval_time: 5m
        alertname: ArgoCDRedisDown
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: "ArgoCD Redis is down"
              description: "ArgoCD Redis instance is down, which may affect the functionality of ArgoCD"
