---

rule_files:
  - ../../resources/prometheus-rules/certificate-expiration.yaml

evaluation_interval: 1h

tests:

  # Test 1
  - interval: 24h
    input_series:
      - series: 'certmanager_certificate_expiration_timestamp_seconds{endpoint="9402",exported_namespace="argocd",instance="10.44.1.34:9402",job="cert-manager-metrics",name="infra-argocd-tls",namespace="infra-cert-manager",pod="cert-manager-6898f8fbcb-6fkx8",service="cert-manager-metrics"}'
        values: '1728000+0x28 4320000+0x30' # expiration = 20 days from start. At day 29, expiration set to day 50

    # Unit test for alerting rules.
    alert_rule_test:

      # No alert :
      - eval_time: 24h
        alertname: CertificateExpired

      # Alert triggered :
      - eval_time: 264h # day 11, 9 days before expiration, throw warning
        alertname: CertificateExpired
        exp_alerts:
        - exp_labels:
            severity: warning
            alertname: CertificateExpired
            endpoint: "9402"
            exported_namespace: argocd
            instance: 10.44.1.34:9402
            job: cert-manager-metrics
            name: infra-argocd-tls
            namespace: infra-cert-manager
            pod: cert-manager-6898f8fbcb-6fkx8
            service: cert-manager-metrics
          exp_annotations:
            summary: '[cert-manager-metrics] Certificate infra-argocd-tls will expire in less than 10 days'
            message: 'Certificate "infra-argocd-tls" in namespace "infra-cert-manager" will expire in 9 days.'


      - eval_time: 384h # day 16, 4 days before expiration, throw critical alert
        alertname: CertificateExpired
        exp_alerts:
        - exp_labels:
            severity: warning
            alertname: CertificateExpired
            endpoint: "9402"
            exported_namespace: argocd
            instance: 10.44.1.34:9402
            job: cert-manager-metrics
            name: infra-argocd-tls
            namespace: infra-cert-manager
            pod: cert-manager-6898f8fbcb-6fkx8
            service: cert-manager-metrics
          exp_annotations:
            summary: '[cert-manager-metrics] Certificate infra-argocd-tls will expire in less than 10 days'
            message: 'Certificate "infra-argocd-tls" in namespace "infra-cert-manager" will expire in 4 days.'
        - exp_labels:
            severity: critical
            alertname: CertificateExpired
            endpoint: "9402"
            exported_namespace: argocd
            instance: 10.44.1.34:9402
            job: cert-manager-metrics
            name: infra-argocd-tls
            namespace: infra-cert-manager
            pod: cert-manager-6898f8fbcb-6fkx8
            service: cert-manager-metrics
          exp_annotations:
            summary: '[cert-manager-metrics] Certificate infra-argocd-tls will expire in less than 5 days'
            message: 'Certificate "infra-argocd-tls" in namespace "infra-cert-manager" will expire in 4 days.'


      - eval_time: 600h # After expiration
        alertname: CertificateExpired
        exp_alerts:
        - exp_labels:
            severity: warning
            alertname: CertificateExpired
            endpoint: "9402"
            exported_namespace: argocd
            instance: 10.44.1.34:9402
            job: cert-manager-metrics
            name: infra-argocd-tls
            namespace: infra-cert-manager
            pod: cert-manager-6898f8fbcb-6fkx8
            service: cert-manager-metrics
          exp_annotations:
            summary: '[cert-manager-metrics] Certificate infra-argocd-tls will expire in less than 10 days'
            message: 'Certificate "infra-argocd-tls" in namespace "infra-cert-manager" will expire in -5 days.'
        - exp_labels:
            severity: critical
            alertname: CertificateExpired
            endpoint: "9402"
            exported_namespace: argocd
            instance: 10.44.1.34:9402
            job: cert-manager-metrics
            name: infra-argocd-tls
            namespace: infra-cert-manager
            pod: cert-manager-6898f8fbcb-6fkx8
            service: cert-manager-metrics
          exp_annotations:
            summary: '[cert-manager-metrics] Certificate infra-argocd-tls will expire in less than 5 days'
            message: 'Certificate "infra-argocd-tls" in namespace "infra-cert-manager" will expire in -5 days.'
        - exp_labels:
            severity: critical
            alertname: CertificateExpired
            endpoint: "9402"
            exported_namespace: argocd
            instance: 10.44.1.34:9402
            job: cert-manager-metrics
            name: infra-argocd-tls
            namespace: infra-cert-manager
            pod: cert-manager-6898f8fbcb-6fkx8
            service: cert-manager-metrics
          exp_annotations:
            summary: '[cert-manager-metrics] Certificate infra-argocd-tls expired'
            message: 'Renew expirated certificate "infra-argocd-tls" in namespace "infra-cert-manager".'


      # at day 30, Alert stopped :
      - eval_time: 720h
        alertname: PrometheusDown


